# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

```python
import pandas as pd

df=pd.read_csv("assets/NISPUF17.csv")
```

Output:
```
Index(['Unnamed: 0', 'SEQNUMC', 'SEQNUMHH', 'PDAT', 'PROVWT_D', 'RDDWT_D',
       'STRATUM', 'YEAR', 'AGECPOXR', 'HAD_CPOX',
       ...
       'XVRCTY2', 'XVRCTY3', 'XVRCTY4', 'XVRCTY5', 'XVRCTY6', 'XVRCTY7',
       'XVRCTY8', 'XVRCTY9', 'INS_STAT2_I', 'INS_BREAK_I'],
      dtype='object', length=454)
```

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    les=df[(df['EDUC1']==1)]
    les1=df[(df['EDUC1']==2)] 
    les2=df[(df['EDUC1']==3)]
    les3=df[(df['EDUC1']==4)]
    total=len(df['EDUC1'])
    a=les['EDUC1'].count()/total
    b=les1['EDUC1'].count()/total
    c=les2['EDUC1'].count()/total
    d=les3['EDUC1'].count()/total
    dicc={"less than high school":a,
    "high school":b,
    "more than high school but not college":c,
    "college":d}
    return dicc
    #raise NotImplementedError()
```

```python
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    milk=df[(df['CBF_01']==1)] 
    r=milk['P_NUMFLU'].dropna().mean()
    milk1=df[(df['CBF_01']==2)]
    r1=milk1['P_NUMFLU'].dropna().mean()
    tupla=(r,r1)
    return tupla
    #raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
DFM=df[(df['SEX']==1)]
DFF=df[(df['SEX']==2)]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D      RDDWT_D  STRATUM  \
0               1   128521     12852     2         NaN   235.916956     1031   
1               2    10741      1074     2         NaN   957.353840     1068   
5               6   207941     20794     2         NaN    35.995367     1019   
8               9   161121     16112     1  978.042448   718.616990     1016   
9              10   143401     14340     2         NaN  1536.859533     1068   
...           ...      ...       ...   ...         ...          ...      ...   
28447       28448   136841     13684     1  593.426753   337.255316     2031   
28448       28449   237881     23788     2         NaN   390.466235     2076   
28449       28450   132081     13208     1  166.681363    96.330007     2059   
28453       28454   122131     12213     2         NaN    33.006992     2013   
28455       28456   154501     15450     2         NaN    56.655283     2062   

       YEAR  AGECPOXR  HAD_CPOX  ...  XVRCTY2 
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex(): 
    MCV=df[(df.P_NUMVRC >=1) & (df.HAD_CPOX == 1) & (df.SEX == 1)]
    MSV=df[(df.P_NUMVRC >=1) & (df.HAD_CPOX == 2) & (df.SEX == 1)]
    FCV=df[(df.P_NUMVRC >=1) & (df.HAD_CPOX == 1) & (df.SEX == 2)]
    FSV=df[(df.P_NUMVRC >=1) & (df.HAD_CPOX == 2) & (df.SEX == 2)]
    M=len(MCV)/len(MSV)
    F=len(FCV)/len(FSV)
    dicc={"male":M,"female":F}
    return dicc
    #raise NotImplementedError()
```

```python
print(chickenpox_by_sex())
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more nos) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isnt really the full picture, since we are not looking at when the dose was given. Its possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
dff=df[["HAD_CPOX","P_NUMVRC"]]
dff=dff[dff.HAD_CPOX.isin([1, 2])]
dff=dff[dff.P_NUMVRC.notnull()]
dff
```

Output:
```
HAD_CPOX  P_NUMVRC
3             2       1.0
4             2       0.0
6             2       1.0
7             2       0.0
8             1       0.0
...         ...       ...
28446         2       1.0
28447         2       1.0
28449         2       1.0
28452         2       1.0
28454         2       1.0

[15286 rows x 2 columns]
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
           #        "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr
    dff=df[["HAD_CPOX","P_NUMVRC"]] 
    dff=dff[dff.HAD_CPOX.isin([1, 2])] 
    dff=dff[dff.P_NUMVRC.notnull()]
    corr, pval=stats.pearsonr(dff["HAD_CPOX"],dff["P_NUMVRC"])
    return corr
    #raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
print(corr_chickenpox())
```

Output:
```
0.07044873460147986
```

```python

```